<div class = "alert alert-warning" align="center"> 
    <img src="https://www.egd.mg/wp-content/uploads/2020/01/cropped-logo-home-110px.png" align="left" width="290"  margin="50px"> 
    <center><h1>Lycée français de Tananarive</h1> 
        <h2> Jupyter Hub,  Spécialité NSI </h2></center>
</div>

<div align=justify>
    <h1 style='text-align:center' ><span style = 'color:#0000FF'> Epréuve Pratique N°4 : Exercice 1</span></h1>
</div>

L'opérateur « ou exclusif » entre deux bits renvoie 0 si les deux bits sont égaux et 1 s'ils sont différents :

0 ⊕ 0 = 0 , 0 ⊕ 1 = 1 , 1 ⊕ 0 = 1 , 1 ⊕ 1 = 0

On représente ici une suite de bits par un tableau contenant des 0 et des 1.
Exemples :
```python
a = [1, 0, 1, 0, 1, 1, 0, 1]
b = [0, 1, 1, 1, 0, 1, 0, 0] 
c = [1, 1, 0, 1]
d = [0, 0, 1, 1]
```

Écrire la fonction xor qui prend en paramètres deux tableaux de même longueur et qui renvoie un tableau où l’élément situé à position i est le résultat, par l’opérateur « ou exclusif », des éléments à la position i des tableaux passés en paramètres.
En considérant les quatre exemples ci-dessus, cette fonction doit passer les tests suivants :
assert(xor(a, b) == [1, 1, 0, 1, 1, 0, 0, 1])
assert(xor(c, d) == [1, 1, 1, 0])

<!-- grade_id: cell-07a62f5a7b4111bd, points: 0, grade: false, solution: true -->
```python
def xor(T1,T2) :
    if len(T1) != len(T2) :
        return None
    resultat = [0 for i in range(len(T1))]
    for j in range(len(T1)) :
        if T1[j] != T2[j] :
            resultat[j] = 1
        else :
            continue
    return resultat

```

<!-- grade_id: cell-e23c11020a8358b7, points: 2, grade: true, solution: false -->
```python
a = [1, 0, 1, 0, 1, 1, 0, 1]
b = [0, 1, 1, 1, 0, 1, 0, 0] 

assert xor(a, b) == [1, 1, 0, 1, 1, 0, 0, 1]

```

<!-- grade_id: cell-03f0ccd8bcd20dcf, points: 2, grade: true, solution: false -->
```python
c = [1, 1, 0, 1]
d = [0, 0, 1, 1]

assert xor(c, d) == [1, 1, 1, 0]
```

<!-- grade_id: cell-bd938c3d0ec1e633, points: 1, grade: true, solution: false -->
```python
a = [1, 0, 1, 0, 1, 1, 0, 1]
c = [1, 1, 0, 1]

assert xor(a, c) == None
```